```python
# Initialisation Cell
from IPython.display import display, HTML, Javascript
from math import *
import numpy as np
import numpy.testing as nt
import time
```

# Numerical Analysis II - Test 2


## Instructions

* You may only access https://moodle.ms.wits.ac.za/tests/ during the test and **NOT** normal moodle (i.e. https: //moodle.ms.wits.ac.za/moodle/ )
* All machines are being logged and anyone caught logging into the regular Moodle listed above will be given **zero**. 
* Anyone caught trying to log into someone else's account will be given **zero**.
* Anyone attempting to use brought in code, notes or any other materials will be given **zero**.
* To launch jupyter notebook correctly. Using VMUSER, open the terminal and type `/usr/local/anaconda3/bin/jupyter-notebook`
* Read all the instructions carefully.
* **DO NOT change the name of the notebook in any way!**
* The test will run from 10:15 - 12:00 (this includes the extra time students). At 12:00 Moodle will no longer take submissions so ensure that you have uploaded your notebook by at least 11:55.
* DO NOT rename the file. Simply answer all questions within the notebook and resubmit the file.
* There are 105 marks available. There are 5 MCQ questions and 5 programming questions.

## MCQ Questions

In each of the MCQ questions there is only one correct answer. Once you have decided an answer in each question, uncomment out the response of your choosing from the corresponding python function.

There are 5 MCQ questions.

***

### Question 1

When implementing a least squares solution, the distance we are minimising is which of the following: [3 marks ] 

* (a) The sum of orthogonal distances between all points and the fit
* (b) The sum of vertical distances between all points and the fit
* (c) The sum of vertical and orthogonal distances between all points and the fit
* (d) There is no minimisation. We actually seek the function which passes through all points
* (e) None of the above


<!-- grade_id: cell-0e0379becdcc85f9, points: 0, grade: false, solution: true -->
```python
def question1():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a47c52014c245c11, points: 0, grade: true, solution: false -->
```python
assert(question1() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-3cb8559d884564aa, points: 3, grade: true, solution: false -->
```python
# Hidden evaulation 
# 3 Marks
```

### Question 2

The family of solutions for the ODE $y^\prime = y$, have what kind of stability: [3 marks ] 

* (a) Stable solutions
* (b) Unstable solutions
* (c) Asymptotically stable solutions
* (d) Insensitive solutions
* (e) None of the above

<!-- grade_id: cell-a3457ca4789a01db, points: 0, grade: false, solution: true -->
```python
def question2():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-359d8a73c61da543, points: 0, grade: true, solution: false -->
```python
assert(question2() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-8445a8c883d76cdf, points: 3, grade: true, solution: false -->
```python
# Hidden evaulation 
# 3 Marks
```

### Question 3

Over time, errors in unstable solutions of ODEs: [3 marks ] 

* (a) Grow with time
* (b) Decrease with time
* (c) Remain the same
* (d) Rapidly decrease with time
* (e) None of the above

<!-- grade_id: cell-9eb61596b73fea32, points: 0, grade: false, solution: true -->
```python
def question3():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-026ca80aa6bd9deb, points: 0, grade: true, solution: false -->
```python
assert(question3() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-029dd73b8907aea0, points: 3, grade: true, solution: false -->
```python
# Hidden evaulation 
# 3 Marks
```

### Question 4

Applying cubic splines to *four* data points yields how many cubic equations? [3 marks ] 

* (a) 4
* (b) 3
* (c) 5
* (d) Depends on the data
* (e) None of the above

<!-- grade_id: cell-868337dee707437d, points: 0, grade: false, solution: true -->
```python
def question4():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-56732a734a441018, points: 0, grade: true, solution: false -->
```python
assert(question4() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-b0f082d681bb67e4, points: 3, grade: true, solution: false -->
```python
# Hidden evaulation 
# 3 Marks
```

### Question 5

Interpolating over an increasingly large number of points may lead to: [3 marks ] 

* (a) A perfectly suitable polynomial
* (b) Gibbs phenomenon
* (c) Runge phenomenon
* (d) (a) or (b)
* (e) None of the above

<!-- grade_id: cell-9f8879b5d88ccc8d, points: 0, grade: false, solution: true -->
```python
def question5():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8b86d1dee5341d2f, points: 0, grade: true, solution: false -->
```python
assert(question5() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-16dbaadd67e23201, points: 3, grade: true, solution: false -->
```python
# Hidden evaulation 
# 3 Marks
```

## Programming Questions 

Complete the following problems:


### Question 1 [15 Marks]

Write a function that will produce a Lagrange interpolating polynomial $P(x)$ and evaluate it at an array of points `xstar`, outputting the resulting `P(xstar)` values. The degree ($n - 1$) polynomial should run through each of the $n$ points ($x, y$) exactly for each point within `xstar`. The inputs `x`, `y` and `xstar` are all numpy arrays, and the corresponding output an array. If the data are not consistent, i.e. `x` and `y` do not have the same length, the function should return the string `'length of data not consistent'`.

**Recall:**
\begin{eqnarray*} L_{n,k}(x)&=&\frac{(x-x_0)\ldots(x-x_{k-1})(x-x_{k+1})\ldots(x-x_n)}{(x_k-x_0)\ldots(x_k-x_{k-1})(x_k-x_{k+1})\ldots(x_k-x_n)}\\
&=&\prod_{i=0,i\neq k}^n \frac{(x-x_i)}{(x_k-x_i)}.\end{eqnarray*}
and:
\begin{equation*}
P(x)=L_{n,0}(x)f(x_0)+L_{n,1}(x)f(x_1)+\ldots +L_{n,n}(x)f(x_n)
\end{equation*}

<!-- grade_id: cell-18e9e616d99d9048, points: 0, grade: false, solution: true -->
```python
def lagrange_polynomial_evaluator(x, y, xstar):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-47eb42c85826bed7, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
xval = np.array([0,1,2,3,4])
yval = np.array([5,11,101,1445])
xstar = np.array([0.5, 1.5, 2.5, 3.5])
nt.assert_string_equal('length of data not consistent', lagrange_polynomial_evaluator(xval, yval, xstar))
print('Test case passed!!!')
```

<!-- grade_id: cell-257d76c448141b18, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
xval = np.array([0,1,2,3,4])
yval = np.array([5,11,101,473,1445])
xstar = np.array([0.5, 1.5, 2.5, 3.5])
nt.assert_array_almost_equal([5.1875, 35.9375, 234.6875, 859.4375], lagrange_polynomial_evaluator(xval, yval, xstar))
print('Test case passed!!!')
```

<!-- grade_id: cell-9ec00b7ed7c85557, points: 12, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 12 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

### Question 2 [15 Marks]

Write a function which performs Euler's method for solving ODEs. The function `euler` should take as inputs, `f` (i.e. $y^\prime = f(x, y)$), the initial value `x0` and `y0`. It should also take the step-size `h` and the final $x$ approximation point `xf`. The function should return an array of the successive $y$ approximations at each step. 


<!-- grade_id: cell-aeb6b321415c5731, points: 0, grade: false, solution: true -->
```python
def euler(f, x0, y0, h, xf):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d291da89848c3757, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
f  = lambda x, y: x + y
h  = 0.1
x0 = 0
y0 = 1
h  = 0.1
xf = 0.4
nt.assert_array_almost_equal([1.1, 1.2200000000000002, 1.362, 1.5282], euler(f, x0, y0, h, xf))
print('All tests passed!!!')
```

<!-- grade_id: cell-5c513679ecaf3157, points: 7, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 7 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

<!-- grade_id: cell-3f7c5e09e709bad0, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

### Question 3 [20 Marks]

Write a function `ODE_solver` which solves a given ODE with both Modified Euler's method and Runge-Kutta 4. The function should take in as input the function `f` (the same format as in Question 2), the initial value `x0` and `y0`, the step-size `h` and the final $x$ value `xf`. The function `ODE_solver` should return a matrix of 2 columns, where column one is the approximation for each $y$ value using the Modified Euler's Method and the column two is approximations using the RK4 method.

**Recall:**


\begin{equation*}
y_{i+1}=y_i+{1\over 6}(k_1+2 k_2+2 k_3+k_4),
\end{equation*}
\begin{eqnarray*}
k_1 &=& hf(x_i,y_i)\\
k_2 &=& hf(x_i+{h\over 2} ,y_i+{k_1\over 2} )\\
k_3 &=& hf(x_i+{h\over 2} ,y_i+{k_2\over 2})\\
k_4 &=& hf(x_i+h ,y_i+k_3),
\end{eqnarray*}
\begin{equation*}
y_{i+\frac{1}{2}}=y_i+f(x_i,y_i) {h\over 2}.
\end{equation*}
\begin{equation*}
y_{i+\frac{1}{2}}^\prime=f(x_{i+1/2},y_{i+1/2}),
\end{equation*}
\begin{equation*}
y_{i+1}=y_i+ f(x_{i+1/2},y_{i+1/2}) h
\end{equation*}

<!-- grade_id: cell-ce0d580c11fa3a11, points: 0, grade: false, solution: true -->
```python
def ODE_solver(f, x0, y0, h, xf):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9d819859fca4faf4, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
y0 = 1
x0 = 0
h  = 0.1
f  = lambda x, y: x + y
xf = 0.4
nt.assert_array_almost_equal(np.array([1.11, 1.24205, 1.39846525, 1.5818041]), ODE_solver(f, x0, y0, h, xf)[:, 0])
print('All tests passed')

```

<!-- grade_id: cell-4c04b2eb6141c917, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
y0 = 1
x0 = 0
h  = 0.1
f  = lambda x, y: x + y
xf = 0.4
nt.assert_array_almost_equal([1.11034167, 1.24280514, 1.39971699, 1.58364848], ODE_solver(f, x0, y0, h, xf)[:, 1])
print('All tests passed')
```

<!-- grade_id: cell-89da06367daa01c3, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
y0 = 1
x0 = 0
h  = 0.1
f  = lambda x, y: x + y
xf = 0.4
tans = [[1.11      , 1.11034167],
       [1.24205   , 1.24280514],
       [1.39846525, 1.39971699],
       [1.5818041 , 1.58364848]]
nt.assert_array_almost_equal(tans, ODE_solver(f, x0, y0, h, xf))
print('All tests passed')

```

<!-- grade_id: cell-5c087891a7a87e22, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

<!-- grade_id: cell-7608ef9c1ce71814, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

<!-- grade_id: cell-8ea7add5c64d2480, points: 6, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

### Question 4 [20 Marks]

Write a function `poly_least_squares` which takes as input some data points `x` ($x$ coordinate) and `y` ($y$ coordinate). It also takes in `order` the degree of polynomial to fit to the given data $(x_i, y_i)$. An `order` value of 1 implies a straight line fit (i.e. $a_0$ and $a_1$). The function should return 3 outputs, the matrix `A`, the vector `b` and the solution to the system `xstar` (the resulting coefficients of the fit).

**Recall**:
\begin{equation*}
f(x)=P_m(x)=a_0+a_1 x+\cdots+a_m x^m
\end{equation*}
\begin{eqnarray*}
a_0 \;n &+& a_1\sum_{i=1}^n x_i+a_2\sum_{i=1}^n x_i^2+\cdots+a_m\sum_{i=1}^n 
x_i^m = \sum_{i=1}^n y_i\\
a_0 \sum_{i=1}^n x_i &+& a_1\sum_{i=1}^n x_i^2+a_2\sum_{i=1}^n x_i^3+\cdots+
a_m\sum_{i=1}^n x_i^{m+1} = \sum_{i=1}^n x_i y_i\\
a_0 \sum_{i=1}^n x_i^2 &+& a_1\sum_{i=1}^n x_i^3+a_2\sum_{i=1}^n x_i^4+\cdots+
a_m\sum_{i=1}^n x_i^{m+2} = \sum_{i=1}^n x_i^2 y_i\\
\vdots && \vdots\\
a_0 \sum_{i=1}^n x_i^m &+& a_1\sum_{i=1}^n x_i^{m+1}+a_2\sum_{i=1}^n x_i^{m+2}+
\cdots+ a_m\sum_{i=1}^n x_i^{2m} = \sum_{i=1}^n x_i^m y_i
\end{eqnarray*}

**Note:** Do not attempt to use any regression or polynomial built-in functions as these will be deleted from the namespace in the hidden tests.

<!-- grade_id: cell-cb2eec749d3ed246, points: 0, grade: false, solution: true -->
```python
def poly_least_squares(x, y, order):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-39333f448853a4ad, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [[  6.,  15.,  55.],
        [ 15.,  55., 225.],
        [ 55., 225., 979.]]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-24ce5f144f77c856, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [ 152.6,  585.6, 2488.8]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-af7a7275f5f4253f, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [2.47857143, 2.35928571, 1.86071429]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[2])
print('Test case passed!!!')
```

<!-- grade_id: cell-0e2f602d34301524, points: 8, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 8 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

<!-- grade_id: cell-73ab1f89971aa917, points: 8, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 8 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

### Question 5 [20 Marks]

Write a function `cubic_splines` which takes input data `x`, `y` and optional `*pt`. The function should apply cubic splines to the data points. Specifically, if no optional input is passed, then the function should return the array's of the coefficients of the cubic splines, i.e. $a_i's$, $b_i's$, $c_i's$ and $d_i's$ (i.e. 4 outputs `a, b, c, d`, where each one is an array). Otherwise, if a point is passed as additional input, then the function should return only the approximation of the appropriate cubic spline for the given point.

**Recall:**

\begin{equation*}
s_i(x) = f_i + b_i(x-x_i) + c_i(x-x_i)^2+d_i(x-x_i)^3.
\end{equation*}
where $a_i = f_i$
\begin{eqnarray*}
b_i &=& \dfrac{f_{i+1}-f_i}{h_i}-\dfrac{h_i}{3}(2c_i + c_{i+1}), \\
d_i &=& \dfrac{c_{i+1}-c_i}{3h_i},
\end{eqnarray*}
\begin{equation*}
h_i = x_{i+1} - x_i.
\end{equation*}
$$
\begin{bmatrix}
1 &  &  & \\ 
h_1 & 2(h_1+h_2)  &h_2  & \\ 
  \ddots &\ddots  &\ddots & \\
 & h_{n-2} & 2(h_{n-2}+h_{n-1})  & h_{n-1} \\
  &  &  & 1 
\end{bmatrix}
\begin{Bmatrix}
c_1\\ 
c_2\\ 
\vdots\\
c_{n-1}\\
c_n 
\end{Bmatrix}
\begin{bmatrix}
0\\ 
3(f[x_3,x_2]-f[x_2,x_1])\\ 
\vdots\\
3(f[x_n,x_{n-1}]-f[x_{n-1},x_{n-2}])\\
0 
\end{bmatrix}.\nonumber
$$

<!-- grade_id: cell-66bab0427a178efd, points: 0, grade: false, solution: true -->
```python
def cubic_splines(x, y, *pt):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0f33ee57f2222e23, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
x  = np.array([3, 4.5, 7, 9])
y  = np.array([2.5, 1, 2.5, 0.5])
t1 = [2.5, 1. , 2.5]
t2 = [-1.41977186, -0.16045627,  0.02205323]
t3 = [ 0.,  0.83954373, -0.76653992,  0.]
t4 = [ 0.18656527, -0.21414449,  0.12775665]
nt.assert_array_almost_equal(t1, cubic_splines(x, y)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-d14ae59e452e55ad, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
x  = np.array([3, 4.5, 7, 9])
y  = np.array([2.5, 1, 2.5, 0.5])
t1 = [2.5, 1. , 2.5]
t2 = [-1.41977186, -0.16045627,  0.02205323]
t3 = [ 0.,  0.83954373, -0.76653992,  0.]
t4 = [ 0.18656527, -0.21414449,  0.12775665]
nt.assert_array_almost_equal(t2, cubic_splines(x, y)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-326e91211567a8c3, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
x  = np.array([3, 4.5, 7, 9])
y  = np.array([2.5, 1, 2.5, 0.5])
t1 = [2.5, 1. , 2.5]
t2 = [-1.41977186, -0.16045627,  0.02205323]
t3 = [ 0.,  0.83954373, -0.76653992,  0.]
t4 = [ 0.18656527, -0.21414449,  0.12775665]
nt.assert_array_almost_equal(t3, cubic_splines(x, y)[2])
print('Test case passed!!!')
```

<!-- grade_id: cell-da2890b4bcfb9d71, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
x  = np.array([3, 4.5, 7, 9])
y  = np.array([2.5, 1, 2.5, 0.5])
t1 = [2.5, 1. , 2.5]
t2 = [-1.41977186, -0.16045627,  0.02205323]
t3 = [ 0.,  0.83954373, -0.76653992,  0.]
t4 = [ 0.18656527, -0.21414449,  0.12775665]
nt.assert_array_almost_equal(t4, cubic_splines(x, y)[3])
print('Test case passed!!!')
```

<!-- grade_id: cell-9dd1111b0166fdcb, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Marks
x = np.array([3, 4.5, 7, 9])
y = np.array([2.5, 1, 2.5, 0.5])
nt.assert_almost_equal(1.102889733840304, cubic_splines(x, y, 5))
print('Test case passed!!!')
```

<!-- grade_id: cell-a2ea808e8d68dd89, points: 8, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 8 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```

<!-- grade_id: cell-1284f1ecde5191c3, points: 7, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 8 Marks
# HIDDEN TEST
# Note nothing will be printed to screen
```