```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
```

# Numerical Analysis II - Lab 1


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

## Warm-up Exercises

Complete the following warm-up tasks:

### Question 1

Write a function that will compute gross pay, i.e. it must take inputs **hours** and **rate** and return **pay**.

<!-- grade_id: cell-26de8dfe31111379, points: 0, grade: false, solution: true -->
```python
def compute_gross_pay(hours, rate):
    """
    Inputs
    hours: scalar input of hours worked
    rate : scalar hourly rate of pay
    
    Outputs
    pay: scalar of total pay earned
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0bbabebbcd5d993f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
x = np.random.randint(10)
y = np.random.randint(150)
nt.assert_approx_equal(x*y, compute_gross_pay(x, y), err_msg = 'function incorrect')
print('All Tests Passed!!!')
```

### Question 2

If we list all the natural number below 10 that are multiples of 3 or 5, we get 3, 5, 6, 9. The sum of these multiples is 23. Write a function that will take in as input **n** and return the sum of all the multiples of 3 or 5 below **n**.

<!-- grade_id: cell-5d3ea102fee7c698, points: 0, grade: false, solution: true -->
```python
def nat_num_multi(n):
    """
    Inputs
    n: integer value
    
    Outputs
    total: sum of all multiples of 3 or 5 below n
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b76f520bfa4e1957, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_approx_equal(23, nat_num_multi(10), err_msg = 'function incorrect')
print('All Tests Passed!!!')
```

### Question 3

Each new term in the Fibonacci sequence is generated by adding the previous two terms. By starting at 1 and 2, the first 10 terms will be:
\begin{equation*}
1, 1, 2, 3, 5, 8, 13, 21, 34, 55, \ldots
\end{equation*}
By considering the terms in the Fibonacci sequence whose values do not exceed input **n**, write a function that returns the sum of the even-valued terms.

<!-- grade_id: cell-1ff8a6ef521af566, points: 0, grade: false, solution: true -->
```python
def fib_sum(n):
    """
    Inputs
    n: integer value
    
    Outputs
    total: sum of the even-valued terms
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5cc3edac6b86680e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_approx_equal(10, 10, err_msg='function incorrect')
print('All Tests Passed!!!')
```

### Question 4

Stirling's approximation is a powerful formula for computing factorials. Write a function to compute the absolute and relative errors in Stirling's approximation:
\begin{equation*}
n! \approx \sqrt{2\pi n}\left(\dfrac{n}{e}\right)^n, 
\end{equation*}
for $n = 1,\ldots,10$. Does the absolute error grow or shrink as $n$ increases? Does the relative error grow or shrink as $n$ increases?

<!-- grade_id: cell-fc8c0cd689060597, points: 0, grade: false, solution: true -->
```python
def stirling_error(n):
    """
    Inputs
    n: integer value
    
    Outputs
    absolute_error, relative_error: absolute error and relative error of Stirling's approximation
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6ca947f02fceee7c, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_array_almost_equal([1.9808320424099009, 0.016506933686749173], 
                             np.asarray(stirling_error(5)), 
                             err_msg = 'function incorrect')
print('All Tests Passed!!!')
```

## Main Exercises


### Question 1

Write functions that can compute the first derivative approximation using backward, forward and central differences.

<!-- grade_id: cell-48d50d22a40f6bf1, points: 0, grade: false, solution: true -->
```python
def fdf(f, h, x):
    """
    Inputs
    f: anonymous function of variable x
    h: scalar step-size
    x: value at which the derivative is approximated at
    
    Outputs
    derivative: the approximation of the first derivative of f(x)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-63147f138ecf1247, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f = lambda x: x**2 - 1
h = 0.1
x = 1
nt.assert_approx_equal(2.100000000000002, fdf(f, h, x))
print('All Tests Passed!!!')
```

<!-- grade_id: cell-977f38f5c4af46ab, points: 0, grade: false, solution: true -->
```python
def bdf(f, h, x):
    """
    Inputs
    f: anonymous function of variable x
    h: scalar step-size
    x: value at which the derivative is approximated at
    
    Outputs
    derivative: the approximation of the first derivative of f(x)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-1eb816e28827a4b6, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_approx_equal(1.8999999999999995, bdf(f, h, x))
print('All Tests Passed!!!')
```

<!-- grade_id: cell-46346acc20f28cf8, points: 0, grade: false, solution: true -->
```python
def cdf(f, h, x):
    """
    Inputs
    f: anonymous function of variable x
    h: scalar step-size
    x: value at which the derivative is approximated at
    
    Outputs
    derivative: the approximation of the first derivative of f(x)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-1f5d246a868c8564, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_approx_equal(2.0000000000000004, cdf(f, h, x))
print('All Tests Passed!!!')
```

### Question 2

Write a function that returns the Richardson's extrapolation table as a matrix. It should use central difference for its initial approximations. The non-entries of the table can be assigned a NaN.

<!-- grade_id: cell-494d1e5949310bae, points: 0, grade: false, solution: true -->
```python
def richardson_table(f, h, x):
    """
    Inputs
    f: anonymous function of variable x
    h: scalar step-size
    x: value at which the derivative is approximated at
    
    Outputs
    r_table: the Richardson's extrapolation table as numpy array
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-3ee7c9beac015dc5, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f = lambda x: x**2*cos(x)
h = np.array([0.1, 0.05, 0.025, 0.0125])
x = 1
tans = [[0.22673616,        nan,        nan,        nan],
       [0.23603092, 0.23912917,        nan,        nan],
       [0.23835774, 0.23913335, 0.23913363,        nan],
       [0.23893964, 0.23913361, 0.23913363, 0.23913363]]
nt.assert_array_almost_equal(tans, richardson_table(f, h, x), err_msg='function incorrect')
print('All Tests Passed!!!')
```