```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
```

# Numerical Analysis II - Lab 6


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* To launch jupyter correctly in the labs using vmuser, type the following at the terminal:
    `/usr/local/anaconda3/bin/jupyter-notebook`

## Main Exercises

Complete the following problems:


### Question 1

Write a function which performs Euler's method for solving ODEs. The function should take as input, $f(x_i, y_i)$ as `f`, the initial values `x0` and `y0` and a stopping value `xf`, with step size `h`. The function should return an approximation to the solution of the ODE at `xf`.

<!-- grade_id: cell-31b9745c4862cdf5, points: 0, grade: false, solution: true -->
```python
def euler(f, x0, y0, h, xf):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-39ee467c2a5442ae, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 3 Marks
f  = lambda x, y: x + y
x0 = 0
y0 = 1
xf = 0.4
h  = 0.1
nt.assert_almost_equal(1.5282, euler(f, x0, y0, h, xf))
print('All tests passed!!!')
```

### Question 2

Write a function which performs Modified Euler's method for solving ODEs. The function should take as input, $f(x_i, y_i)$ as `f`, the initial values `x0` and `y0` and a stopping value `xf`, with step size `h`. The function should return an approximation to the solution of the ODE at `xf`.

<!-- grade_id: cell-20bdb73160560497, points: 0, grade: false, solution: true -->
```python
def modified_euler(f, x0, y0, h, xf):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2dbac53c8509d81f, points: 4, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 4 Marks
f  = lambda x, y: x + y
x0 = 0
y0 = 1
xf = 0.4
h  = 0.1
nt.assert_array_almost_equal(1.5818041012500001, modified_euler(f, x0, y0, h, xf))
print('All tests passed!!!')
```

### Question 3

Write a function which performs Runge-Kutta 4th order for solving ODEs. The function should take as input, $f(x_i, y_i)$ as `f`, the initial values `x0` and `y0` and a stopping value `xf`, with step size `h`. The function should return an approximation to the solution of the ODE at `xf`.

<!-- grade_id: cell-53d4f0e214f52065, points: 0, grade: false, solution: true -->
```python
def rk4(f, x0, y0, h, xf):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-abe7d85b38b2a6b6, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
f  = lambda x, y: x + y
x0 = 0
y0 = 1
xf = 0.4
h  = 0.1
nt.assert_almost_equal(1.5836484801613715, rk4(f, x0, y0, h, xf))
print('All tests passed!!!')
```

<!-- grade_id: cell-0cd6ace5407e6908, points: 8, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 8 Marks
# Hidden Tests
```