```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
```

# Numerical Analysis II - Lab 5


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* To launch jupyter correctly in the labs using vmuser, type the following at the terminal:
    `/usr/local/anaconda3/bin/jupyter-notebook`

## Main Exercises

Complete the following problems:


### Question 1

Write a function `poly_least_squares` which takes as input some data points `x` ($x$ coordinate) and `y` ($y$ coordinate). It also takes in `order` the degree of polynomial to fit to the given data $(x_i, y_i)$. An `order` value of 1 implies a straight line fit (i.e. $a_0$ and $a_1$). The function should return 3 outputs, the matrix `A`, the vector `b` and the solution to the system `xstar` (the resulting coefficients of the fit).

<!-- grade_id: cell-df97b2184a6d6040, points: 0, grade: false, solution: true -->
```python
def poly_least_squares(x, y, order):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-70f05b82f802ee42, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [[  6.,  15.,  55.],
        [ 15.,  55., 225.],
        [ 55., 225., 979.]]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-fa1bc5f23207e3fe, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [ 152.6,  585.6, 2488.8]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-a6338611f6bb7ecb, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
x = np.array([0, 1, 2, 3, 4, 5])
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
tans = [2.47857143, 2.35928571, 1.86071429]
nt.assert_array_almost_equal(tans, poly_least_squares(x, y, 2)[2])
print('Test case passed!!!')
```

<!-- grade_id: cell-666af35fa061e163, points: 10, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 10 Marks
x  = np.array([1, 2, 3, 4, 5, 6, 7])
y  = np.array([0.5, 2.5, 2.0, 4.0, 3.5, 6.0, 5.5])
t1 = np.array([[  7.,  28. ],
               [ 28.,  140. ]])
t2 = np.array([ 24. , 119.5])
t3 = np.array([0.07142857, 0.83928571])
nt.assert_almost_equal(t1, (poly_least_squares(x, y, 1)[0]))
nt.assert_almost_equal(t2, (poly_least_squares(x, y, 1)[1]))
nt.assert_almost_equal(t3, (poly_least_squares(x, y, 1)[2]))
print('Test case passed!!!')
```